```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

Output:
```
c:\Users\toema\miniforge3\envs\dsc80\Lib\site-packages\otter\export\__init__.py:6: UserWarning: pkg_resources is deprecated as an API. See https://setuptools.pypa.io/en/latest/pkg_resources.html. The pkg_resources package is slated for removal as early as 2025-11-30. Refrain from using this package or pin to Setuptools<81.
  import pkg_resources
```

# Lab 6 – APIs and Web Scraping

## DSC 80, Spring 2026

### Due Date: Wednesday, May 13th at 11:59 PM


## Instructions

Welcome to the sixth DSC 80 lab this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2026-wi).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

<div class="alert alert-block alert-danger">
<b>Note: </b> For this lab, due to system constraints, the code may take varying amounts of time to run on different machines. If you are able to submit to Gradescope and the code runs, then you have met the runtime requirements.
</div>

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

If you encounter a `ModuleNotFoundError`, please run `!pip install lxml` in a new cell. After `lxml` is succesfully installed, go to `kernel` then restart. Note that you will only need to run `!pip install lxml` once. 

```python
!pip install lxml
```

Output:
```
Requirement already satisfied: lxml in c:\users\toema\miniforge3\envs\dsc80\lib\site-packages (6.1.0)
```

```python
import os
import pandas as pd
import numpy as np
np.set_printoptions(legacy='1.21')
import requests
import bs4
import lxml
```

```python
#!pip install lxml
```

## Question 1 – Practice with HTML Tags 📎

In Question 2, you'll spend plenty of time parsing HTML source code. But before you get your hands dirty trying to extract information from HTML written by other people, it is a good idea to write basic HTML code yourself. This exercise will help you better understand how the code in a `.html` file is structured.

For this question, you'll create a very basic `.html` file, named `lab06_1.html`, that satisfies the following conditions:

- It must have `<title>` and `<head>` tags.
- It must also have `<body>` tags. Within the `<body>` tags, it must have:
    - At least two headers.
    * At least three images.
        - At least one image must be a local file.
        - At least one image must be linked to an online source.
        - At least one image has to have default text when it cannot be displayed.
    * At least three references (hyperlinks) to different web pages.
    * At least one table with two or more rows and two or more columns.
    

Make sure to save your file as `lab06_1.html`, and save it in the same directory as `lab.py`. **When submitting this homework to Gradescope, make sure to also upload `lab06_1.html` along with the local image that you embedded in your site.** You can upload multiple files to Gradescope at a time.
   

***Notes***:
- You can write and view basic HTML with a Jupyter Notebook, using either a Markdown cell or by using the `IPython.display.HTML` function (which takes in a string of HTML and renders it).
- If you write your HTML code within a Jupyter Notebook, you should later copy your code into a text editor and save it with the `.html` extension. You could also write your HTML in a text editor directly.
- Be sure to open your final `.html` file in a browser and make sure it looks correct on its own.

```python
def question1():
    """Returns the path to lab06_1.html (the file must exist)."""
    return 'lab06_1.html'

```

```python

```

```python
# Don't delete this cell!
question1()
```

Output:
```
'lab06_1.html'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Question 2 – Scraping an Online Bookstore 📚

Browse through the following fake online bookstore: http://books.toscrape.com/. This website is meant for toying with scraping.

Your job is to scrape the website, collecting data on all books that:
- **have _at least_ a four-star rating**, and
- **cost _strictly_ less than £50**, and 
- **belong to specific categories** (more details below). 

You will extract the information into a DataFrame that looks like the one below.

<table border="1" class="dataframe">
  <thead>
    <tr style="text-align: right;">
      <th></th>
      <th>UPC</th>
      <th>Product Type</th>
      <th>Price (excl. tax)</th>
      <th>Price (incl. tax)</th>
      <th>Tax</th>
      <th>Availability</th>
      <th>Number of reviews</th>
      <th>Category</th>
      <th>Rating</th>
      <th>Description</th>
      <th>Title</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <th>0</th>
      <td>e10e1e165dc8be4a</td>
      <td>Books</td>
      <td>Â£22.60</td>
      <td>Â£22.60</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Default</td>
      <td>Four</td>
      <td>For readers of Laura Hillenbrand's Seabiscuit...</td>
      <td>The Boys in the Boat: Nine Americans...</td>
    </tr>
    <tr>
      <th>1</th>
      <td>c2e46a2ee3b4a322</td>
      <td>Books</td>
      <td>Â£25.27</td>
      <td>Â£25.27</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Romance</td>
      <td>Five</td>
      <td>A Michelin two-star chef at twenty-eight, Violette...</td>
      <td>Chase Me (Paris Nights #2)</td>
    </tr>
    <tr>
      <th>2</th>
      <td>00bfed9e18bb36f3</td>
      <td>Books</td>
      <td>Â£34.53</td>
      <td>Â£34.53</td>
      <td>Â£0.00</td>
      <td>In stock (19 available)</td>
      <td>0</td>
      <td>Romance</td>
      <td>Five</td>
      <td>No matter how busy he keeps himself...</td>
      <td>Black Dust</td>
    </tr>
  </tbody>
</table>

To do so, implement the following functions.

<br>

#### `extract_book_links`

Complete the implementation of the function `extract_book_links`, which takes in the content of a page that contains book listings as a **string of HTML**, and returns a **list** of URLs of book-specific pages for all books with **_at least_ a four-star rating and a price _strictly_ less than £50**.

For this method, the URLs you return should not contain the protocol (i.e. `'https://'`). The protocols should be added back into the URLs when you actually make the requests.


<br>

#### `get_product_info`

Complete the implementation of the function `get_product_info`, which takes in the content of a book-specific page as a **string of HTML**, and a list `categories` of book categories. Note that each book belongs to exactly one category. If the input book's category is in the list of `categories`, `get_product_info` should return a dictionary corresponding to a row in the DataFrame in the image above (where the keys are the column names and the values are the row values). If the input book's category is not in the list of `categories`, return `None`. <b> The order of the dictionary keys does not matter. </b>


<br>

#### `scrape_books`

Finally, put everything together. Complete the implementation of the function `scrape_books`, which takes in an integer `k` and a list `categories` of book categories. `scrape_books` should use `requests` to scrape the first `k` pages of the bookstore and return a DataFrame of only the books that have:
- **_at least_ a four-star rating**, and
- **a price _strictly_ less than £50**, and
- **a category that is in the list `categories`**.

<b> The order of the DataFrame columns does not matter. </b>

<br>

Some general guidance and tips:

- The first page of the bookstore is at http://books.toscrape.com/catalogue/page-1.html. Subsequent pages can be found by clicking the "Next" button at the bottom of the page. Look at how the URLs change each time you navigate to a new page; think about how to use f-strings (or some other string formatting technique) to generate these URLs.
- Use "inspect element" to view the source code of the pages you're trying to scrape. To find a book's category, look at the hyperlinks in the book-specific page for that book.
- **`scrape_books` should run in under 180 seconds on the entire bookstore (`k = 50`). `scrape_books` is also the only function that should make `GET` requests; the other two functions parse already-existing HTML.**
- When instantiating `bs4.BeautifulSoup` objects, use the optional argument `features='lxml'`.
- Don't worry about typecasting, i.e. it's fine if `'Number of reviews'` is not stored as type `int`. Also, don't worry if you run into encoding errors in your price columns (as the example DataFrame at the top of this cell contains).

```python
RATING_MAP = {'One': 1, 'Two': 2, 'Three': 3, 'Four': 4, 'Five': 5}
 
 
def extract_book_links(page_html):
    """
    Given HTML of a catalogue listing page (as a string), return a list of
    relative URLs for books that have at least a 4-star rating AND a price
    strictly less than £50.  URLs should NOT contain the 'https://' protocol.
    """
    soup = bs4.BeautifulSoup(page_html, features='lxml')
    links = []
 
    for article in soup.find_all('article', class_='product_pod'):
        # Rating
        rating_tag = article.find('p', class_='star-rating')
        rating_word = rating_tag['class'][1]  # e.g. 'Four'
        rating = RATING_MAP.get(rating_word, 0)
        if rating < 4:
            continue
 
        # Price (strip the £ sign and any encoding artifacts)
        price_text = article.find('p', class_='price_color').get_text(strip=True)
        # Remove currency symbol (could be encoded oddly) – keep digits and '.'
        price_str = ''.join(c for c in price_text if c.isdigit() or c == '.')
        price = float(price_str)
        if price >= 50:
            continue
 
        # URL
        a_tag = article.find('h3').find('a')
        href = a_tag['href']  # relative URL like 'scarlet.../index.html'
        links.append(href)
 
    return links
 
 
def get_product_info(page_html, categories):
    """
    Given HTML of a book-specific page (as a string) and a list of acceptable
    category names, return a dict of product info if the book's category is in
    `categories`, otherwise return None.
    """
    soup = bs4.BeautifulSoup(page_html, features='lxml')
 
    # Category is the last breadcrumb link before the book title
    breadcrumbs = soup.find('ul', class_='breadcrumb').find_all('li')
    # breadcrumbs: Home > Books > <Category> > <Title>
    category = breadcrumbs[-2].get_text(strip=True)
 
    if category not in categories:
        return None
 
    # Rating
    rating_tag = soup.find('p', class_='star-rating')
    rating = rating_tag['class'][1]  # word form e.g. 'Two'
 
    # Title
    title = soup.find('div', class_='product_main').find('h1').get_text(strip=True)
 
    # Description – the first <p> after the product description div
    desc_div = soup.find('div', id='product_description')
    if desc_div:
        description = desc_div.find_next_sibling('p').get_text(strip=True)
    else:
        description = ''
 
    # Product info table
    table = soup.find('table', class_='table table-striped')
    info = {}
    for row in table.find_all('tr'):
        key = row.find('th').get_text(strip=True)
        val = row.find('td').get_text(strip=True)
        info[key] = val
 
    result = {
        'UPC': info.get('UPC', ''),
        'Product Type': info.get('Product Type', ''),
        'Price (excl. tax)': info.get('Price (excl. tax)', ''),
        'Price (incl. tax)': info.get('Price (incl. tax)', ''),
        'Tax': info.get('Tax', ''),
        'Availability': info.get('Availability', ''),
        'Number of reviews': info.get('Number of reviews', ''),
        'Category': category,
        'Rating': rating,
        'Description': description,
        'Title': title,
    }
    return result
 
 
def scrape_books(k, categories):
    """
    Scrape the first k pages of books.toscrape.com and return a DataFrame of
    books that have ≥4-star rating, price < £50, and category in `categories`.
    """
    base_url = 'http://books.toscrape.com/catalogue/page-{}.html'
    book_base = 'http://books.toscrape.com/catalogue/'
 
    rows = []
    for page_num in range(1, k + 1):
        page_url = base_url.format(page_num)
        resp = requests.get(page_url)
        page_html = resp.text
 
        links = extract_book_links(page_html)
        for link in links:
            # links from the listing page are relative like 'book-title/index.html'
            book_url = book_base + link
            book_resp = requests.get(book_url)
            info = get_product_info(book_resp.text, categories)
            if info is not None:
                rows.append(info)
 
    if rows:
        return pd.DataFrame(rows)
    else:
        return pd.DataFrame(columns=[
            'UPC', 'Product Type', 'Price (excl. tax)', 'Price (incl. tax)',
            'Tax', 'Availability', 'Number of reviews', 'Category',
            'Rating', 'Description', 'Title'
        ])

```

```python

```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests

# public test for extract_book_links 
extract_book_links_fp = os.path.join('data', 'products.html')
extract_book_out = extract_book_links(
    open(extract_book_links_fp, encoding='utf-8').read()
)
extract_book_url = 'scarlet-the-lunar-chronicles-2_218/index.html'

# doc tests for get product info
get_product_info_fp = os.path.join('data', 'Frankenstein.html')
get_product_info_out = get_product_info(
    open(get_product_info_fp, encoding='utf-8').read(), ['Default']
)

# public test for scrape books 
scrape_books_out = scrape_books(1, ['Mystery'])
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Question 3 – API Requests 🤑

You trade stocks as a hobby. As an avid `pandas` coder, you decide to calculate statistics of your favorite stocks by pulling data from a public API. The API we will work with can be found at https://financialmodelingprep.com/developer/docs. Specifically, we will use the **Stock Price and Volume Data API** endpoint (search for it at the linked page).

Some relevant definitions:
- Symbol: A short code that refers to a stock. For example, Apple's symbol is AAPL and Ford's symobl is F. 
- Open: The price of a stock at the beginning of a trading day.
- Close: The price of a stock at the end of a trading day.
- Volume: The total number of shares traded in a day.
- Percent change: The difference in price with respect to the original price, as a percentage.

To make requests to the aforementioned API, you will need an API key. In order to get one, you will need to make an account at the website. Once you've signed up, you can use the API key that comes with the free plan. It has a limit of 250 requests per day, which should be more than enough. You will have to encode your API key in the URL that you make requests to; see the Authentication section at the top of the [documentation](https://site.financialmodelingprep.com/developer/docs).

Implement the following two functions.

#### `stock_history`

Complete the implementation of the function `stock_history`, which takes in a string `symbol` and two integers, `year` and `month`, and returns a DataFrame containing the price history for that stock in that month. Keep all of the attributes that are returned by the API.

***Notes***:
- Read the API documentation to learn how to request data for a specific date range.
- [`pd.date_range`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.date_range.html) takes in two dates and returns a sequence of all dates between the two dates, excluding the right endpoint. This may be helpful!
- The [`requests.get`](https://docs.python-requests.org/en/master/user/quickstart/) function returns a Response object, not the data itself. Use the `json` method on the Response object to extract the relevant JSON, as we did in [Lecture 9](https://dsc80.com/resources/lectures/lec09/lec09.html#Example:-GET-requests-via-requests) (you don't need to `import json` to do this).

<br>

#### `stock_stats`

Create a function `stock_stats` that takes in a DataFrame outputted by `stock_history` and returns a **tuple** of two numbers:
1. The percent change of the stock throughout the month as a **percentage**.
2. The total transaction volume **in billion of dollars** for that month.

Both values in the tuple should be **strings** representing numbers rounded two decimal places. Add a plus or minus sign in front of the percent change and the percent symbol at the end. Make sure that the total transaction volume string ends in a `'B'` for billions. There should always be two decimal places specified (for example `'35.40%'`, not `'35.4%'`).

**To compute the percent change**, use the opening price on the first day of the month as the starting price and the closing price on the last day of the month as the ending price.

**To compute the total transaction volume**, assume that on any given day, the average price of a share is the average of the high and low prices for that day. The daily transaction volume is given by the formula below, and the total transaction volume is the sum of the daily transaction volumes across all days of the month.

$$ \text{Daily Transaction Volume (in dollars)} = \text{Volume (number of shares traded)} \times \text{Average Price} $$

For example, suppose there are only three days in March – March 1st, March 2nd, and March 3rd.

If BYND (Beyond Meat) opens at \\$4 on March 1st and closes at \\$5 on March 3rd, its percent change for the month of March is $$\frac{\$5-\$4}{\$4} = +25.00\%$$

Suppose the high and low prices and volumes of BYND on each day are given below.
- March 1st: high \\$5, low \\$3, volume 500 million (0.5 billion)
- March 2nd: high \\$5.5, low \\$2.5, volume 1 billion
- March 3rd: high \\$5.25, low \\$4, volume 500 million (0.5 billion)

Then, the estimated total transaction volume is
$$\frac{\$5 + \$3}{2} \cdot 0.5 B + \frac{\$5.5 + \$2.5}{2} \cdot 1 B + \frac{\$5.25 + \$4}{2} \cdot 0.5 B = 8.3125B$$

```python
FMP_API_KEY = 'nQ9vhyRrSfCfTrHL8dDtHlmWZ3wfpJpJ'
 
def stock_history(symbol, year, month):
    """
    Return a DataFrame of daily price history for `symbol` during the given
    year/month, using the Financial Modeling Prep API.
    """
    start = pd.Timestamp(year=year, month=month, day=1)
    end = start + pd.offsets.MonthEnd(1)
    from_date = start.strftime('%Y-%m-%d')
    to_date = end.strftime('%Y-%m-%d')
 
    url = (
        f'https://financialmodelingprep.com/stable/historical-price-eod/full'
        f'?symbol={symbol}&from={from_date}&to={to_date}&apikey={FMP_API_KEY}'
    )
    resp = requests.get(url)
    data = resp.json()
 
    if isinstance(data, list):
        df = pd.DataFrame(data)
    elif isinstance(data, dict) and 'historical' in data:
        df = pd.DataFrame(data['historical'])
    else:
        raise KeyError(f"Unexpected API response: {data}")
 
    return df
 
 
def stock_stats(df):
    """
    Given a stock history DataFrame, return a tuple of two strings:
    1. Percent change (e.g. '+3.45%' or '-1.23%')
    2. Total transaction volume in billions (e.g. '8.31B')
    """
    # Sort ascending by date so first row = earliest day
    df = df.sort_values('date', ascending=True).reset_index(drop=True)
 
    start_price = df.iloc[0]['open']
    end_price = df.iloc[-1]['close']
    pct_change = (end_price - start_price) / start_price * 100
 
    sign = '+' if pct_change >= 0 else ''
    pct_str = f'{sign}{pct_change:.2f}%'
 
    # Total transaction volume
    df['avg_price'] = (df['high'] + df['low']) / 2
    df['daily_volume_dollars'] = df['volume'] * df['avg_price']
    total_volume_billions = df['daily_volume_dollars'].sum() / 1e9
    vol_str = f'{total_volume_billions:.2f}B'
 
    return (pct_str, vol_str)

```

```python

```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests

# public test for stock_history
history = stock_history('AAPL', 2019, 2)

# public test for stock_stats
stats = stock_stats(history)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Question 4 – Comment Threads 🧵

You regularly browse [Hacker News](https://news.ycombinator.com/) to keep up with the latest news in tech. One link to a Hacker News article is https://news.ycombinator.com/item?id=18344932. Note that this article has 18 comments and has a `storyid` of 18344932.

The problem now is that you don't have internet access on your phone during your morning commute to work, so you want to save the interesting stories' comment threads beforehand in a CSV. You find their [API documentation](https://github.com/HackerNews/API) and decide to get to work.

Complete the implementation of the function `get_comments`, which takes in a `storyid` and returns a DataFrame of all the comments below the news story. You can ignore "dead" comments (you will know them when you see them), as well as "dead" comments’ children. **Make sure the order of the comments in your DataFrame is from top to bottom just as you see on the website**. 

The DataFrame that `get_comments` returns should have 5 columns:
1. `'id'`: The unique ID of the comment.
2. `'by'`: The author of the comment.
3. `'text'`: The actual comment.
4. `'parent'`: The unique ID of the comment this comment is replying to.
5. `'time'`: When the comment was created (in `pd.Timestamp` format).

Some guidance:
1. The URL to make requests to is `'https://hacker-news.firebaseio.com/v0/item/{}.json'`, however, the `{}` should be replaced with the ID of the article or page you are trying to access. 
2. Again, do not `import json` – instead, use the `json` method on the Response object you get back.
3. Use depth-first search when traversing the comments tree. You will have to do this manually, since you cannot use Beautiful Soup (which is only for HTML documents, not JSON objects).
4. Make sure the length of your returned DataFrame is the same as value for the `'descendants'` key in the response JSON (both of which correspond to the number of comments for the story).
5. You are allowed to use loops in this function. You may also want to create at least one helper function.

<div class="alert alert-block alert-success">
    You may find <a href="https://www.youtube.com/watch?v=uOfwW-onmpc"><b>this hint video 🎥</b></a> helpful!
</div>

```python
 
HN_URL = 'https://hacker-news.firebaseio.com/v0/item/{}.json'
 
 
def _fetch_item(item_id):
    resp = requests.get(HN_URL.format(item_id))
    return resp.json()
 
 
def _collect_comments(item_id, rows):
    """Depth-first traversal collecting non-dead comments."""
    item = _fetch_item(item_id)
    if item is None:
        return
    # Skip dead comments
    if item.get('dead', False):
        return
    # Skip deleted comments (they may lack 'by'/'text')
    if item.get('deleted', False):
        return
 
    row = {
        'id': item.get('id'),
        'by': item.get('by', ''),
        'text': item.get('text', ''),
        'parent': item.get('parent'),
        'time': pd.Timestamp(item.get('time', 0), unit='s'),
    }
    rows.append(row)
 
    for child_id in item.get('kids', []):
        _collect_comments(child_id, rows)
 
 
def get_comments(storyid):
    """
    Return a DataFrame of all comments for the given Hacker News story,
    traversed depth-first (top to bottom as on the website).
    """
    story = _fetch_item(storyid)
    rows = []
    for kid_id in story.get('kids', []):
        _collect_comments(kid_id, rows)
 
    df = pd.DataFrame(rows, columns=['id', 'by', 'text', 'parent', 'time'])
    return df.reset_index(drop=True)

```

```python

```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests
comments = get_comments(18344932)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Congratulations! You're done with Lab 6! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q2 q4
```

will run the `grader.check` cells for Questions 1, 2, and 4 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!
```

```python

```

Tests for q1 (points: None)
```
>>> os.path.exists('lab06_1.html')
True
```
```
>>> q01_html = bs4.BeautifulSoup(open("lab06_1.html", encoding='utf-8'), features='lxml')
>>> len(q01_html.find_all('image')) == 0
True
```
failure_message: you should not use the <image> tag to add images


Tests for q2 (points: None)
```
>>> extract_book_out[1] == extract_book_url
True
```
```
>>> isinstance(get_product_info_out, dict)
True
```
```
>>> 'Category' in get_product_info_out.keys()
True
```
```
>>> get_product_info_out['Rating'] == 'Two'
True
```
```
>>> scrape_books_out.shape == (1, 11)
True
```
```
>>> scrape_books_out['Rating'][0] == 'Four'
True
```
```
>>> scrape_books_out['Title'][0] == 'Sharp Objects'
True
```


Tests for q3 (points: None)
```
>>> history.shape == (19, 10)
True
```
```
>>> history.date.iloc[-1] == '2019-02-01'
True
```
```
>>> (len(stats[0]), len(stats[1])) == (6, 6)
True
```
```
>>> (float(stats[0][1:-1]) < 10) == True
True
```
```
>>> float(stats[1][:-1]) > 1 == True
True
```
```
>>> stats[1][-1] == 'B'
True
```


Tests for q4 (points: None)
```
>>> comments.shape == (18, 5)
True
```
```
>>> comments.loc[5, 'by'] == 'RobAtticus'
True
```
```
>>> comments.loc[5, 'time'].day == 31
True
```
